```python
# Run this cell first!
import otter
import inspect
import sys
from io import StringIO
import ast
from IPython.display import display, Javascript
grader = otter.Notebook("Homework2.ipynb")

def capture_output():
    # Backup the original stdout
    original_stdout = sys.stdout

    # Create a StringIO object to capture the output
    sys.stdout = StringIO()

    return original_stdout

def restore_output(original_stdout):
    # Capture the output and restore the original stdout
    captured_output = sys.stdout.getvalue()
    
    sys.stdout = original_stdout
    
    return captured_output
```

Output:
```
ModuleNotFoundError                       Traceback (most recent call last)
Cell In[1], line 2
      1 # Run this cell first!
----> 2 import otter
      3 import inspect
      4 import sys

ModuleNotFoundError: No module named 'otter'
```

# Data Programming in Python | BAIS:6040
# Homework 2. Python Basics Part 2

10 questions, 6 points in total

## Instructions

- Fill in the Python statements needed to complete all the questions below. 
- Program everything inside this notebook. 
- Use only one cell per question. 
- Each code cell must be run with no errors.
- Do not contain any unnecessary code in your answer. (Only the final outcome will be regarded as your answer.)
- If the question requests that you store information in a certain variable, use the specific variables name (case sensitive).
- All problems must be solved in a programmable way.
- When you are done, make sure to save your notebook and then click Restart & Run All before subimission to create a clean version.

## Questions

1\. (0.6 pts) You have a string `s`. Write a <b>for</b> loop that appends each letter in `s` in reverse order with no space between letters. The printed output should look as follows:

zyxwvutsrqponmlkjihgfedcba

```python
# Run this cell. Do not modify anything in this cell.
s = "abcdefghijklmnopqrstuvwxyz"
```

```python
for letter in reversed(s):
    print(letter, end="")
```

Output:
```
zyxwvutsrqponmlkjihgfedcba
```

```python

    
```

```python
# DO NOT MODIFY THE NEXT LINE
original_stdout = capture_output()

# Your answer starts here
for letter in reversed(s):
    print(letter, end="")

# DO NOT MODIFY ANYTHING BELOW THIS LINE
print_capture_q1 = restore_output(original_stdout)
print(print_capture_q1)
```

```python
grader.check("Question01")
```

<!-- BEGIN QUESTION -->

2\. (0.6 pts) Write a <b>for</b> loop that prints all multiples of 3 between 1 and 100 (i.e., 3, 6, 9, ..., 99), separated by a tab between numbers. The output should match the next cell.

3	6	9	12	15	18	21	24	27	30	33	36	39	42	45	48	51	54	57	60	63	66	69	72	75	78	81	84	87	90	93	96	99

```python
# DO NOT MODIFY THE NEXT LINE
original_stdout = capture_output()

# Your answer starts here
...

# DO NOT MODIFY ANYTHING BELOW THIS LINE
print_capture_q2 = restore_output(original_stdout)
print(print_capture_q2)
```

```python
grader.check("Question02")
```

<!-- END QUESTION -->

3\. (0.6 pts) Write a <b>for</b> loop that adds all 2-digit integers (i.e., 10, 11, 12, ..., 99) one after another. After the <b>for</b> loop, print the final sum, which will be 4905.

```python
# DO NOT MODIFY THE NEXT LINE
original_stdout = capture_output()

# Your answer starts here
...

# DO NOT MODIFY ANYTHING BELOW THIS LINE
print_capture_q3 = restore_output(original_stdout)
print(print_capture_q3)
```

```python
grader.check("Question03")
```

4\. (0.6 pts) Write a <b>for</b> loop that adds all 2-digit integers one after another, printing the intermediate sum for every 10th index position starting with the first position (e.g., 0, 10, 20, ..., 80), separated by a new line, in the format "[INDEX] SUM" where INDEX is the current index position and SUM is the current sum. Print the final sum after the <b>for</b> loop. The first and last three lines in the printed output should look follows:

[0] 10
[10] 165
[20] 420

[70] 3195
[80] 4050
4905

```python
# DO NOT MODIFY THE NEXT LINE
original_stdout = capture_output()

# Your answer here
...

# DO NOT MODIFY ANYTHING BELOW THIS LINE
print_capture_q4 = restore_output(original_stdout)
print(print_capture_q4)
```

```python
grader.check("Question04")
```

5\. (0.6 pts) You have a list `names` of strings. 

```python
# Make sure to run this cell. (Do not make any change to this cell.)
names = ["Alice", "Bob", "Mary", "Sam", "Sarah", "Tom", "Susan"]
```

Write a <b>for</b> loop that iterates over `names` to find a name that meets the following two conditions: a) it starts with *S* and b) its length is 5. The loop stops when it finds the first occurrence of the name, printing the name found along with its index position in the format "[INDEX] NAME", where INDEX is the current index position and NAME is the name found. The printed output should look as follows:

[4] Sarah

```python
# DO NOT MODIFY THE NEXT LINE
original_stdout = capture_output()

# Your answer here
...

# DO NOT MODIFY ANYTHING BELOW THIS LINE
print_capture_q5 = restore_output(original_stdout)
print(print_capture_q5)
```

```python
grader.check("Question05")
```

6\. (0.6 pts) You have a dictionary `states`. 

```python
# Make sure to run this cell. (Do not make any change to this cell.)
states = {"IL": "Illinois", "WI": "Wisconsin", "IA": "Iowa", "NE": "Nebraska", "MN": "Minnesota"}
```

Write a <b>for</b> loop that prints all key-value pairs in `states`, separated by a space between pairs, in the format "(KEY: VALUE)", where KEY is the key and VALUE is its value. The printed output should look as follows: 

(IL: Illinois) (WI: Wisconsin) (IA: Iowa) (NE: Nebraska) (MN: Minnesota)

```python
# DO NOT MODIFY THE NEXT LINE
original_stdout = capture_output()

# Your answer starts here
...


# DO NOT MODIFY ANYTHING BELOW THIS LINE
print_capture_q6 = restore_output(original_stdout)
print(print_capture_q6)
```

```python
grader.check("Question06")
```

7\. (0.6 pts) You have two lists `l1` and `l2` of integers. 

```python
# Make sure to run this cell. (Do not make any change to this cell.)
l1 = list(range(5))
l2 = list(reversed(l1))
l1, l2
```

Using a basic nested <b>for</b> loop, not using list comprehension, create a list `ans7` of all possible lists, each of which contains a number from `l1`, a number from `l2`, and the product of the two numbers. For example, `ans7` should contain the elements as follows:

[[0, 4, 0], 
 [0, 3, 0], 
 [0, 2, 0], 
 [0, 1, 0], 
 [0, 0, 0], 
 [1, 4, 4], 
 ..., 
 [4, 0, 0]]

Note: Do not use the <b>print</b> function in this question. This question does not ask you to print anything.

```python
# DO NOT MODIFY THE NEXT LINE
original_stdout = capture_output()

# Your answer here
...

# DO NOT MODIFY ANYTHING BELOW THIS LINE
print(ans7)
print_capture_q7 = restore_output(original_stdout)
print(print_capture_q7)
```

```python
grader.check("Question07")
```

8\. (0.6 pts) Using list comprehension this time, create a list `ans8` that has the same elements as `ans7`.

Note: Do not use the <b>print</b> function in this question. This question does not ask you to print anything. 

```python
# Your answer starts here
ans8 = ...
```

```python
grader.check("Question08")
```

<!-- BEGIN QUESTION -->

9\. (0.6 pts) You have three lists `names`, `ages`, and `genders`. 

```python
# Make sure to run this cell. (Do not make any change to this cell.)
names = ["Alice", "Bob", "Mary", "Sam", "Sarah", "Tom"]
ages = [30, 25, 27, 31, 39, 34]
genders = ["female", "male", "female", "male", "female", "male"] 
```

Using a basic <b>for</b> loop, not using list comprehension, create a list `ans9` of tuples, each of which contains the name from `names`, the age from `ages`, and the gender from `genders` at the same index position. For example, `ans9` should contain the elements as follows:

[('Alice', 30, 'female'),
 ('Bob', 25, 'male'), 
 ..., 
 ('Tom', 34, 'male')]

Note: Do not use the <b>print</b> function in this question. This question does not ask you to print anything. 

```python
# Your answer here
...
```

```python
# Check your answer here. (Do not make any change to this cell. Just run this cell.)
ans9
```

```python
grader.check("Question09")
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

10\. (0.6 pts) Using list comprehension this time, create a list `ans10` that has the same elements as `ans9`.

Note: do not use the <b>print</b> function in this question. This question does not ask you to print anything. 

```python
# Your answer here
ans10 = ...
```

```python
# Check your answer here. (Do not make any change to this cell. Just run this cell.)
ans10
```

```python
grader.check("Question10")
```

<!-- END QUESTION -->

### Now that you have completed all the questions, click Kernel > Restart & Run All Cells to and check if all cells have run as expected. 

## Quick Survey

How long did it take you to complete the homework?

: X minutes

Tests for Question01 (points: 0.6)
```
>>> def test_solution_q1():
...     try:
...         assert print_capture_q1 == 'zyxwvutsrqponmlkjihgfedcba', 'The output is incorrect.'
...         print('Test passed!')
...     except AssertionError as e:
...         print(e)
>>> test_solution_q1()
Test passed!
```


Tests for Question02 (points: 0.6)
```
>>> def test_solution_q2():
...     expected_output = '\t'.join((str(num) for num in range(3, 100, 3))) + '\t'
...     try:
...         assert print_capture_q2 == expected_output, 'The output is incorrect.'
...         print('Test passed!')
...     except AssertionError as e:
...         print(e)
>>> test_solution_q2()
Test passed!
```


Tests for Question03 (points: 0.6)
```
>>> def test_solution_q3():
...     expected_output = '4905\n'
...     try:
...         assert print_capture_q3 == expected_output, 'The output is incorrect.'
...         print('Test passed!')
...     except AssertionError as e:
...         print(e)
>>> test_solution_q3()
Test passed!
```


Tests for Question04 (points: 0.6)
```
>>> def test_solution_q4():
...     expected_output = '[0] 10\n[10] 165\n[20] 420\n[30] 775\n[40] 1230\n[50] 1785\n[60] 2440\n[70] 3195\n[80] 4050\n4905\n'
...     try:
...         assert print_capture_q4 == expected_output, 'The output is incorrect.'
...         print('Test passed!')
...     except AssertionError as e:
...         print(e)
>>> test_solution_q4()
Test passed!
```


Tests for Question05 (points: 0.6)
```
>>> def test_solution_q5():
...     expected_output = '[4] Sarah\n'
...     try:
...         assert print_capture_q5 == expected_output, 'The output is incorrect.'
...         print('Test passed!')
...     except AssertionError as e:
...         print(e)
>>> test_solution_q5()
Test passed!
```


Tests for Question06 (points: 0.6)
```
>>> def test_solution_q6():
...     expected_output = '(IL: Illinois) (WI: Wisconsin) (IA: Iowa) (NE: Nebraska) (MN: Minnesota) '
...     try:
...         assert print_capture_q6 == expected_output, 'The output is incorrect.'
...         print('Test passed!')
...     except AssertionError as e:
...         print(e)
>>> test_solution_q6()
Test passed!
```


Tests for Question07 (points: 0.6)
```
>>> def test_solution_q7():
...     expected_output = '[[0, 4, 0], [0, 3, 0], [0, 2, 0], [0, 1, 0], [0, 0, 0], [1, 4, 4], [1, 3, 3], [1, 2, 2], [1, 1, 1], [1, 0, 0], [2, 4, 8], [2, 3, 6], [2, 2, 4], [2, 1, 2], [2, 0, 0], [3, 4, 12], [3, 3, 9], [3, 2, 6], [3, 1, 3], [3, 0, 0], [4, 4, 16], [4, 3, 12], [4, 2, 8], [4, 1, 4], [4, 0, 0]]\n'
...     try:
...         assert print_capture_q7 == expected_output, 'The output is incorrect.'
...         print('Test passed!')
...     except AssertionError as e:
...         print(e)
>>> test_solution_q7()
Test passed!
```


Tests for Question08 (points: 0.6)
```
>>> def test_solution_q8():
...     expected_output = [[0, 4, 0], [0, 3, 0], [0, 2, 0], [0, 1, 0], [0, 0, 0], [1, 4, 4], [1, 3, 3], [1, 2, 2], [1, 1, 1], [1, 0, 0], [2, 4, 8], [2, 3, 6], [2, 2, 4], [2, 1, 2], [2, 0, 0], [3, 4, 12], [3, 3, 9], [3, 2, 6], [3, 1, 3], [3, 0, 0], [4, 4, 16], [4, 3, 12], [4, 2, 8], [4, 1, 4], [4, 0, 0]]
...     try:
...         assert ans8 == expected_output, 'The output is incorrect.'
...         print('Test passed!')
...     except AssertionError as e:
...         print(e)
>>> test_solution_q8()
Test passed!
```


Tests for Question09 (points: 0.6)
```
>>> def test_solution_q9():
...     expected_output = [('Alice', 30, 'female'), ('Bob', 25, 'male'), ('Mary', 27, 'female'), ('Sam', 31, 'male'), ('Sarah', 39, 'female'), ('Tom', 34, 'male')]
...     try:
...         assert ans9 == expected_output, 'The output is incorrect.'
...         print('Test passed!')
...     except AssertionError as e:
...         print(e)
>>> test_solution_q9()
Test passed!
```


Tests for Question10 (points: 0.6)
```
>>> def test_solution_q10():
...     expected_output = [('Alice', 30, 'female'), ('Bob', 25, 'male'), ('Mary', 27, 'female'), ('Sam', 31, 'male'), ('Sarah', 39, 'female'), ('Tom', 34, 'male')]
...     try:
...         assert ans10 == expected_output, 'The output is incorrect.'
...         print('Test passed!')
...     except AssertionError as e:
...         print(e)
>>> test_solution_q10()
Test passed!
```
