```python
# Run this cell first!
import otter
grader = otter.Notebook("Homework3.ipynb")
```

```python
#pip install otter-grader
```

# Data Programming in Python | BAIS:6040
# Homework 3. User-Defined Functions, Exceptions, and Numpy

10 questions, 6 points in total

## Instructions

- Fill in the Python statements needed to complete all the questions below. 
- Program everything inside this notebook. 
- Use only one cell per question. 
- Each code cell must be run with no errors.
- Do not contain any unnecessary code in your answer. (Only the final outcome will be regarded as your answer.)
- If the question requests that you store information in a certain variable, use the specific variables name (case sensitive).
- All problems must be solved in a programmable way.
- When you are done, make sure to save your notebook and then click Restart & Run All before subimission to create a clean version.

## Questions

1\. (0.6 pts) Write a function `get_stats` that takes a list `l` of integers as an argument and returns a tuple of the min, max, and sum of only the positive numbers in `l` for each. You may need a temporary list like `l_tmp` to store only the positive numbers in `l`. Use a for loop to fill in `l_tmp`.

Note: Do not use the <b>print</b> function in this question. The function returns a tuple, not prints it. 

```python
# Your answer here
def get_stats(l):
    l_tmp = [num for num in l if num > 0]
    if not l_tmp:
        return (None, None, 0)
    return (min(l_tmp), max(l_tmp), sum(l_tmp))
```

```python
# Test your function here. (Do not make any change to this cell. Just run this cell.)
get_stats([1, 3, 2])            # Output: (1, 3, 6)
```

Output:
```
(1, 3, 6)
```

```python
# Test your function here. (Do not make any change to this cell. Just run this cell.)
get_stats([-1, 1, 0, 3, 2])     # Output: (1, 3, 6)
```

Output:
```
(1, 3, 6)
```

```python
# Test your function here. (Do not make any change to this cell. Just run this cell.)
get_stats([5, 10, 15])     # Output: (5, 15, 30)
```

Output:
```
(5, 15, 30)
```

```python
grader.check("Question01")
```

Output:
```
Question01 results: All test cases passed!
```

2\. (0.6 pts) Write a function `get_stats2` that works the same as `get_stats` above, but, additionally, it returns the  '*Invalid value!*' string when it finds any non-integer in `l`. Use a for loop to check if each item in `l` is an integer.

Note: Do not use the <b>print</b> function in this question. The function returns either a tuple or a string, not prints it. 

```python
# Your answer here
def get_stats2(l):
    for item in l:
        if not isinstance(item, int):
            return 'Invalid value!'
    l_tmp = [num for num in l if num > 0]
    if not l_tmp:
        return (None, None, 0)
    return (min(l_tmp), max(l_tmp), sum(l_tmp))
```

```python
# Test your function here. (Do not make any change to this cell. Just run this cell.)
get_stats2([-1, 1, 0, 3, 2])     # Output: (1, 3, 6)
```

Output:
```
(1, 3, 6)
```

```python
# Test your function here. (Do not make any change to this cell. Just run this cell.)
get_stats2([-1, 1, 0.0, 3, 2])   # Output: 'Invalid value!'.
```

Output:
```
'Invalid value!'
```

```python
# Test your function here. (Do not make any change to this cell. Just run this cell.)
get_stats2([-1, 1, "0", 3, 2])   # Output: 'Invalid value!'.
```

Output:
```
'Invalid value!'
```

```python
grader.check("Question02")
```

Output:
```
Question02 results: All test cases passed!
```

3\. (0.6 pts) Write a function `truncate` that takes a string `s` and an integer `limit` as arguments and returns the string as it is if the length of the string is equal to or less than `limit`. If the length is greater than `limit`, it only takes the `limit`-length leading substring of `s` (i.e., from the index position 0 to `limit` -1) replacing the rest of `s` with '...' and returns that truncated string. Set the default value for `limit` to 50.

Note: Do not use the <b>print</b> function in this question. The function returns a string, not prints anything. 

```python
# Your answer here
def truncate(s, limit=50):
    if len(s) <= limit:
        return s
    return s[:limit] + '...'
```

```python
# Test your function here. (Do not make any change to this cell. Just run this cell.)
truncate("I'm learning Python data analytics.", 100)

# Output: "I'm learning Python data analytics."
```

Output:
```
"I'm learning Python data analytics."
```

```python
# Test your function here. (Do not make any change to this cell. Just run this cell.)
truncate("I'm learning Python data analytics.", 10)

# Output: "I'm learni..."
```

Output:
```
"I'm learni..."
```

```python
# Test your function here. (Do not make any change to this cell. Just run this cell.)
truncate("I'm learning Python data analytics. Python is a general purpose programming language.")

# Output: "I'm learning Python data analytics. Python is a ge..."
```

Output:
```
"I'm learning Python data analytics. Python is a ge..."
```

```python
grader.check("Question03")
```

Output:
```
Question03 results: All test cases passed!
```

4\. (0.6 pts) Create a dictionary `kwargs` that sets the argument values of `s` to any string longer than 30 characters and `limit` to 30 and then call the `truncate` function by passing `kwargs`.

```python
# Your answer here 

kwargs = {
    's': "I'm learning Python data analytics",
    'limit': 30
}

# Call the truncate function with the dictionary
truncate(**kwargs)
```

Output:
```
"I'm learning Python data analy..."
```

```python
grader.check("Question04")
```

Output:
```
Question04 results: All test cases passed!
```

5\. (0.6 pts) To use NumPy, import the NumPy package as np.

```python


# Your answer here
import numpy as np
```

```python
grader.check("Question05")
```

Output:
```
Question05 results: All test cases passed!
```

6\. (0.6 pts) Set the random seed to 0 and then create a 10 x 10 array `x` of random integers between 1 (inclusive) and 100 (inclusive).

```python
# Your answer here 
np.random.seed(0)
x = np.random.randint(1, 101, size=(10, 10))
```

```python
# Check your answer here. (Do not make any change to this cell. Just run this cell.)
x
```

Output:
```
array([[ 45,  48,  65,  68,  68,  10,  84,  22,  37,  88],
       [ 71,  89,  89,  13,  59,  66,  40,  88,  47,  89],
       [ 82,  38,  26,  78,  73,  10,  21,  81,  70,  80],
       [ 48,  65,  83, 100,  89,  50,  30,  20,  20,  15],
       [ 40,  33,  66,  10,  58,  33,  32,  75,  24,  36],
       [ 76,  56,  29,  35,   1,   1,  37,  54,   6,  39],
       [ 18,  80,   5,  43,  59,  32,   2,  66,  42,  58],
       [ 36,  12,  47,  83,  92,   1,  15, 100,  54,  13],
       [ 43,  85,  76,  69,   7,  69,  48,   4,  77,  53],
       [ 79,  16,  21, 100,  59,  24,  80,  14,  86,  49]], dtype=int32)
```

```python
grader.check("Question06")
```

Output:
```
Question06 results: All test cases passed!
```

7\. (0.6 pts) Create a 10 x 10 array `y` of integers from 100 (inclusive) to 1 (inclusive), such that the first row has integers from 100 to 91, the second row from 90 to 81, and so on and so forth. Note that order of values matters. 

```python
# Your answer here

y = np.arange(100, 0, -1).reshape(10, 10)
```

```python
# Check your answer here. (Do not make any change to this cell. Just run this cell.)
y
```

Output:
```
array([[100,  99,  98,  97,  96,  95,  94,  93,  92,  91],
       [ 90,  89,  88,  87,  86,  85,  84,  83,  82,  81],
       [ 80,  79,  78,  77,  76,  75,  74,  73,  72,  71],
       [ 70,  69,  68,  67,  66,  65,  64,  63,  62,  61],
       [ 60,  59,  58,  57,  56,  55,  54,  53,  52,  51],
       [ 50,  49,  48,  47,  46,  45,  44,  43,  42,  41],
       [ 40,  39,  38,  37,  36,  35,  34,  33,  32,  31],
       [ 30,  29,  28,  27,  26,  25,  24,  23,  22,  21],
       [ 20,  19,  18,  17,  16,  15,  14,  13,  12,  11],
       [ 10,   9,   8,   7,   6,   5,   4,   3,   2,   1]])
```

```python
grader.check("Question07")
```

Output:
```
Question07 results: All test cases passed!
```

8\. (0.6 pts) Select all the values that are in the first 3 rows and the last 7 columns of `x` and save the resulting array in `x_top_right`. Select all the values that are in the last 7 rows and the first 3 columns of `y` and save the resulting array in `y_bottom_left`.

```python
# Your answer here 
x_top_right = x[:3, -7:]
y_bottom_left = y[-7:, :3]
```

```python
# Check your answer here. (Do not make any change to this cell. Just run this cell.)
x_top_right
```

Output:
```
array([[68, 68, 10, 84, 22, 37, 88],
       [13, 59, 66, 40, 88, 47, 89],
       [78, 73, 10, 21, 81, 70, 80]], dtype=int32)
```

```python
# Check your answer here. (Do not make any change to this cell. Just run this cell.)
y_bottom_left
```

Output:
```
array([[70, 69, 68],
       [60, 59, 58],
       [50, 49, 48],
       [40, 39, 38],
       [30, 29, 28],
       [20, 19, 18],
       [10,  9,  8]])
```

```python
grader.check("Question08")
```

Output:
```
Question08 results: All test cases passed!
```

9\. (0.6 pts) Before we calculate the dot product of `x_top_right` and `y_bottom_left`, we want to confirm that the number of columns of `x_top_right` equals the number of rows of `y_bottom_left`. Compare the two values using the equality operator (==) and save the resulting True value in `ans9`. Do not directly put their values (i.e., 7) in the comparison.

```python
# Your answer here 
ans9 = x_top_right.shape[1] == y_bottom_left.shape[0]
```

```python
# Check your answer here. (Do not make any change to this cell. Just run this cell.)
ans9
```

Output:
```
True
```

```python
grader.check("Question09")
```

Output:
```
Question09 results: All test cases passed!
```

10\. (0.6 pts) Calculate the dot product of `x_top_right` and `y_bottom_left` and then save the resulting array in the variable `dot_prod`. 

```python
# Your answer here 
dot_prod = np.dot(x_top_right, y_bottom_left)
```

```python
# Check your answer here. (Do not make any change to this cell. Just run this cell.)
dot_prod
```

Output:
```
array([[14980, 14603, 14226],
       [13820, 13418, 13016],
       [15810, 15397, 14984]])
```

```python
grader.check("Question10")
```

Output:
```
Question10 results: All test cases passed!
```

### Now that you have completed all the questions, click Kernel > Restart & Run All to creat a clean version of this notebook and check if all cells have run as expected. 

Tests for Question01 (points: 0.6)
```
>>> def test_get_stats():
...     assert get_stats([1, 3, 2]) == (1, 3, 6), 'Test case 1 failed'
...     assert get_stats([-1, 1, 0, 3, 2]) == (1, 3, 6), 'Test case 2 failed'
...     assert get_stats([5, 10, 15]) == (5, 15, 30), 'Test case 3 failed'
...     print('All test cases for get_stats passed!')
>>> test_get_stats()
All test cases for get_stats passed!
```


Tests for Question02 (points: 0.6)
```
>>> def test_get_stats2():
...     assert get_stats2([-1, 1, 0, 3, 2]) == (1, 3, 6), 'Test case 1 failed'
...     assert get_stats2([-1, 1, 0.0, 3, 2]) == 'Invalid value!', 'Test case 2 failed'
...     assert get_stats2([-1, 1, '0', 3, 2]) == 'Invalid value!', 'Test case 3 failed'
...     print('All test cases for get_stats2 passed!')
>>> test_get_stats2()
All test cases for get_stats2 passed!
```


Tests for Question03 (points: 0.6)
```
>>> def test_truncate():
...     assert truncate("I'm learning Python data analytics.", 100) == "I'm learning Python data analytics.", 'Test case 1 failed'
...     assert truncate("I'm learning Python data analytics.", 10) == "I'm learni...", 'Test case 2 failed'
...     assert truncate('Short text', 50) == 'Short text', 'Test case 3 failed'
...     assert truncate('This is a very long string that should be truncated.', 20) == 'This is a very long ...', 'Test case 4 failed'
...     assert truncate('Exact length string.', 20) == 'Exact length string.', 'Test case 5 failed'
...     assert truncate("I'm learning Python data analytics. Python is a general purpose programming language.") == "I'm learning Python data analytics. Python is a ge...", 'Test case 6 failed'
...     print('All test cases for truncate passed!')
>>> test_truncate()
All test cases for truncate passed!
```


Tests for Question04 (points: 0.6)
```
>>> def test_truncate_kwargs():
...     try:
...         assert 'kwargs' in globals(), "The variable 'kwargs' is not defined."
...     except AssertionError as e:
...         print(e)
...         return
...     assert 's' in kwargs, "The key 's' is not in kwargs."
...     assert len(kwargs['s']) > 30, "The value for key 's' is not longer than 30 characters."
...     assert 'limit' in kwargs, "The key 'limit' is not in kwargs."
...     assert kwargs['limit'] == 30, "The value for key 'limit' is not 30."
...     s = kwargs['s']
...     limit = kwargs['limit']
...     if len(s) <= limit:
...         expected_output = s
...     else:
...         expected_output = s[:limit] + '...'
...     actual_output = truncate(**kwargs)
...     print('Actual Output:', actual_output)
...     print('Expected Output:', expected_output)
...     assert actual_output == expected_output, 'The output of truncate(**kwargs) is incorrect.'
...     print('All test cases for truncate with kwargs passed!')
>>> test_truncate_kwargs()
Actual Output: I'm learning Python data analy...
Expected Output: I'm learning Python data analy...
All test cases for truncate with kwargs passed!
```


Tests for Question05 (points: 0.6)
```
>>> def test_import_numpy():
...     try:
...         assert 'np' in globals(), "The package 'numpy' is not imported as 'np'."
...     except AssertionError as e:
...         print(e)
...         return
...     try:
...         import numpy
...         assert np is numpy, "'np' is not the numpy package."
...     except AssertionError as e:
...         print(e)
...         return
...     print('All test cases for importing numpy passed!')
>>> test_import_numpy()
All test cases for importing numpy passed!
```


Tests for Question06 (points: 0.6)
```
>>> def test_random_array():
...     np.random.seed(0)
...     expected_x = np.random.randint(1, 101, (10, 10))
...     try:
...         assert 'x' in globals(), "The variable 'x' is not defined."
...     except AssertionError as e:
...         print(e)
...         return
...     assert x.shape == (10, 10), "The shape of 'x' is not (10, 10)."
...     assert np.all((x >= 1) & (x <= 100)), "The values in 'x' are not all between 1 and 100 (inclusive)."
...     assert np.array_equal(x, expected_x), "The values in 'x' do not match the expected values."
...     print('All test cases for the random array passed!')
>>> test_random_array()
All test cases for the random array passed!
```


Tests for Question07 (points: 0.6)
```
>>> def test_array_y():
...     try:
...         assert 'y' in globals(), "The variable 'y' is not defined."
...     except AssertionError as e:
...         print(e)
...         return
...     assert y.shape == (10, 10), "The shape of 'y' is not (10, 10)."
...     expected_y = np.arange(100, 0, -1).reshape(10, 10)
...     assert np.array_equal(y, expected_y), "The values in 'y' do not match the expected values."
...     print("All test cases for the array 'y' passed!")
>>> test_array_y()
All test cases for the array 'y' passed!
```


Tests for Question08 (points: 0.6)
```
>>> def test_subarrays():
...     try:
...         assert 'x_top_right' in globals(), "The variable 'x_top_right' is not defined."
...     except AssertionError as e:
...         print(e)
...         return
...     try:
...         assert 'y_bottom_left' in globals(), "The variable 'y_bottom_left' is not defined."
...     except AssertionError as e:
...         print(e)
...         return
...     expected_x_top_right = x[:3, -7:]
...     assert np.array_equal(x_top_right, expected_x_top_right), "The values in 'x_top_right' are not correct."
...     expected_y_bottom_left = y[-7:, :3]
...     assert np.array_equal(y_bottom_left, expected_y_bottom_left), "The values in 'y_bottom_left' are not correct."
...     print('All test cases for the subarrays passed!')
>>> test_subarrays()
All test cases for the subarrays passed!
```


Tests for Question09 (points: 0.6)
```
>>> def test_ans9():
...     try:
...         assert 'ans9' in globals(), "The variable 'ans9' is not defined."
...     except AssertionError as e:
...         print(e)
...         return
...     assert isinstance(ans9, bool), "The variable 'ans9' is not a boolean."
...     expected_ans9 = x_top_right.shape[1] == y_bottom_left.shape[0]
...     assert ans9 == expected_ans9, "The value of 'ans9' is not correct."
...     print('All test cases for ans9 passed!')
>>> test_ans9()
All test cases for ans9 passed!
```


Tests for Question10 (points: 0.6)
```
>>> def test_dot_product():
...     try:
...         assert 'dot_prod' in globals(), "The variable 'dot_prod' is not defined."
...     except AssertionError as e:
...         print(e)
...         return
...     expected_dot_prod = np.dot(x_top_right, y_bottom_left)
...     assert np.array_equal(dot_prod, expected_dot_prod), "The values in 'dot_prod' are not correct."
...     print('All test cases for dot_prod passed!')
>>> test_dot_product()
All test cases for dot_prod passed!
```
