```python
# Initialize Otter
import otter
grader = otter.Notebook("Homework1.ipynb")
```

# Data Programming in Python | BAIS:6040
# Homework 1. Python Basics Part 1

20 questions, 3 points in total

## Instructions

- Fill in the Python statements needed to complete all the questions below. 
- Program everything inside this notebook. 
- Use only one cell per question. 
- Each code cell must be run with no errors.
- Do not contain any unnecessary code in your answer. (Only the final outcome will be regarded as your answer.)
- If the question requests that you store information in a certain variable, use the specific variables name (case sensitive).
- All problems must be solved in a programmable way.
- When you are done, make sure to save your notebook and then click Restart & Run All before subimission to create a clean version.

### Important
If you haven't installed the `otter` package, please follow these steps:
1. <b>Uncomment</b> the line below by reomving the `#` at the beginning.
2. <b>Run</b> the cell to install the package. You only need to do this once.
3. After installation, <b>re-comment</b> the line or delete to prevent it from running again in future executions of this notebook.

```python
#pip install otter-grader
```

## Questions

Suppose you have a string `s` as below. Do not change the value of `s` while working on the questions. 

```python
# Make sure to run this cell. (Do not make any change to this cell.)
s = "I'm taking the Data Programming in Python course in this semester."
s
```

Output:
```
"I'm taking the Data Programming in Python course in this semester."
```

1\. (0.1 pts) Get the length of `s` and save it in the variable `ans1`. Remember to assign an object, e.g., an integer, to a variable you need to use the assignment operator, "=", e.g., `ans1 = myanswer`.

```python
# Your answer here
ans1=len(s)
```

```python
grader.check("Question01")
```

Output:
```
Question01 results: All test cases passed!
```

2\. (0.1 pts) Get the first 3 characters in `s` and save it in the variable `ans2`.

```python
# Your answer here 
ans2=s[:3]
```

```python
grader.check("Question02")
```

Output:
```
Question02 results: All test cases passed!
```

3\. (0.1 pts) Get the last 10 characters in `s` and save it in the variable `ans3`.

```python
# Your answer here
ans3=s[-10:]
```

```python
grader.check("Question03")
```

Output:
```
Question03 results: All test cases passed!
```

4\. (0.1 pts) Check if the string '*Python*' is in `s` and save the outcome in the variable `ans4`. The outcome must be either True or False.

```python
# Check your answer here. (Do not make any change to this cell. Just run this cell.)
ans4="Python" in s 
```

```python
grader.check("Question04")
```

Output:
```
Question04 results: All test cases passed!
```

5\. (0.1 pts) Replace '*Data Programming in Python*' in `s` with '*Text Analytics*' and save the resulting string in the variable `ans5`. The new string must look like this: "*I'm taking the Text Analytics course in this semester.*"

```python
# Your answer here
ans5=s.replace("Data Programming in Python", "Text Analytics")
```

```python
grader.check("Question05")
```

Output:
```
Question05 results: All test cases passed!
```

6\. (0.1 pts) Get the frequency of the string '*the*' in `s` and save it in the variable `ans6`.

```python
# Your answer here
ans6=s.count("the")
```

```python
grader.check("Question06")
```

Output:
```
Question06 results: All test cases passed!
```

<hr>

Suppose you have a list `heroes` as follows:

```python
# Make sure to run this cell.
heroes = ["Superman", "Ant-Man", "Spider-Man", "Iron Man", "Bat Man"]
```

7\. (0.1 pts) Append a new item '*Aquaman*' to `heroes`. You do not have to save the outcome in a variable. 

```python
# Your answer here

heroes.append("Aquaman")
```

```python
grader.check("Question07")
```

Output:
```
Question07 results: All test cases passed!
```

8\. (0.1 pts) Insert a new item '*Captain America*' at the first index position (i.e., index 0) in `heroes`. You do not have to save the outcome in a variable. 

```python
# Your answer here

newentry='Captain America'
heroes.insert(0,newentry )
```

```python
grader.check("Question08")
```

Output:
```
Question08 results: All test cases passed!
```

9\. (0.1 pts) Get the number of items in `heroes` and save it in the variable `ans9`.

```python
# Your answer here
ans9=len(heroes)
```

```python
grader.check("Question09")
```

Output:
```
Question09 results: All test cases passed!
```

10\. (0.1 pts) Get a new sorted list of `heroes` in reverse alphabetical order and save it in the variable `ans10`. Do not change the value of `heroes`.

```python
# Your answer here
ans10=sorted(heroes, reverse=True)
```

```python
grader.check("Question10")
```

Output:
```
Question10 results: All test cases passed!
```

<hr>

Suppose you have a dictionary `states` as follows.

```python
# Make sure to run this cell
states = {"Iowa": "IA", "Illinois": "IL", "Minnesota": "MN", "Nebraska": "NE", "Wisconsin": "WI"}
```

11\. (0.2 pts) Get the abbreviation for '*Iowa*' using `states` and save it in the variable `ans11`.

```python
# Your answer here
ans11 = states["Iowa"]
```

```python
grader.check("Question11")
```

Output:
```
Question11 results: All test cases passed!
```

12\. (0.2 pts) Check if the key '*Missouri*' is in `states` and save the outcome in the variable `ans12`. The outcome must be either True or False.

```python
# Your answer here
ans12="Missouri" in states
```

```python
grader.check("Question12")
```

Output:
```
Question12 results: All test cases passed!
```

13\. (0.2 pts) Add '*Missouri*' and its abbreviation '*MO*' to `states`. You do not have to save the outcome in a variable. 

```python
# Your answer here

states["Missouri"] = "MO"
```

```python
grader.check("Question13")
```

Output:
```
Question13 results: All test cases passed!
```

14\. (0.2 pts) Get the list of all keys (only the keys, not the key-value pairs) in `states` and save it in the variable `ans14`.

```python
# Your answer here
ans14 = list(states.keys())
```

```python
grader.check("Question14")
```

Output:
```
Question14 results: All test cases passed!
```

<hr>

Suppose you have a list `grades` as follows. 

```python
# Make sure to run this cell
grades = ["A", "C", "A", "B", "B", "C", "A", "B", "A", "A"]
```

15\. (0.2 pts) Convert the list `grades` to a set and save the resulting set in the variable `grades_set`. Do not change the value of `grades` for later use.

```python
# Your answer here
grades_set = set(grades)
```

```python
grader.check("Question15")
```

Output:
```
Question15 results: All test cases passed!
```

16\. (0.2 pts) Get the number of unique items in `grades` and save it in the variable `ans16`. Use `grades_set`.

```python
# Your answer here

ans16 = len(grades_set)
```

```python
grader.check("Question16")
```

Output:
```
Question16 results: All test cases passed!
```

17\. (0.2 pts) Check if the length of `grades` is greater than or equal to the length of `grades_set` and save the outcome in the variable `ans17`. The outcome must be either True or False.

```python
# Your answer here
ans17 = len(grades) >= len(grades_set)
```

```python
grader.check("Question17")
```

Output:
```
Question17 results: All test cases passed!
```

<hr>

18\. (0.2 pts) Calculate the remainder of 283 divided by 10 and save it in the variable `ans18`.

```python
# Your answer here
ans18 = 283 % 10
```

```python
grader.check("Question18")
```

Output:
```
Question18 results: All test cases passed!
```

19\. (0.2 pts) Calculate the value of 5 to the power of 4 and save it in the variable `ans19`.

```python
# Your answer here
ans19 = 5 ** 4
```

```python
grader.check("Question19")
```

Output:
```
Question19 results: All test cases passed!
```

20\. (0.2 pts) Delete the variable `grades`. You do not have to save the outcome in a variable. 

```python
# Your answer here
del grades
```

```python
grader.check("Question20")
```

Output:
```
Question20 results: All test cases passed!
```

## Quick Survey

How long did it take you to complete this homework? (This will not affect your score. Your answer will be very helpful for the instructor to design the homewok assignments and exams. Thank you!)

: 15 minutes

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Output:
```
Running your submission against local test cases...


Your submission received the following results when run against available test cases:

    Question01 results: All test cases passed!

    Question02 results: All test cases passed!

    Question03 results: All test cases passed!

    Question04 results: All test cases passed!

    Question05 results: All test cases passed!

    Question06 results: All test cases passed!

    Question07 results: All test cases passed!

    Question08 results: All test cases passed!

    Question09 results: All test cases passed!

    Question10 results: All test cases passed!

    Question11 results: All test cases passed!

    Question12 results: All test cases passed!

    Question13 results: All test cases passed!

    Question14 results: All test cases passed!

    Question15 results: All test cases passed!

    Question16 results: All test cases passed!

    Question17 results: All test cases passed!

    Question18 results: All test cases pass
```

Tests for Question01 (points: 0.1)
```
>>> def test_q1():
...     try:
...         assert ans1 == 66
...         print('Test passed!')
...     except AssertionError as e:
...         print(e)
>>> test_q1()
Test passed!
```


Tests for Question02 (points: 0.1)
```
>>> def test_q2():
...     try:
...         assert ans2 == "I'm"
...         print('Test passed!')
...     except AssertionError as e:
...         print(e)
>>> test_q2()
Test passed!
```


Tests for Question03 (points: 0.1)
```
>>> def test_q3():
...     try:
...         assert ans3 == ' semester.'
...         print('Test passed!')
...     except AssertionError as e:
...         print(e)
>>> test_q3()
Test passed!
```


Tests for Question04 (points: 0.1)
```
>>> def test_q4():
...     try:
...         assert ans4 == True
...         print('Test passed!')
...     except AssertionError as e:
...         print(e)
>>> test_q4()
Test passed!
```
```
>>> ans4
True
```


Tests for Question05 (points: 0.1)
```
>>> def test_q5():
...     try:
...         assert ans5 == "I'm taking the Text Analytics course in this semester.", 'The replacement string is incorrect.'
...         print('Test passed!')
...     except AssertionError as e:
...         print(e)
>>> test_q5()
Test passed!
```


Tests for Question06 (points: 0.1)
```
>>> def test_q6():
...     try:
...         assert ans6 == 1, "The frequency of 'the' is incorrect."
...         print('Test passed!')
...     except AssertionError as e:
...         print(e)
>>> test_q6()
Test passed!
```


Tests for Question07 (points: 0.1)
```
>>> def test_q7():
...     try:
...         assert heroes[-1] == 'Aquaman', "The item 'Aquaman' was not appended correctly."
...         print('Test passed!')
...     except AssertionError as e:
...         print(e)
>>> test_q7()
Test passed!
```


Tests for Question08 (points: 0.1)
```
>>> def test_q8():
...     try:
...         assert heroes[0] == 'Captain America', "The item 'Captain America' was not inserted correctly."
...         print('Test passed!')
...     except AssertionError as e:
...         print(e)
>>> test_q8()
Test passed!
```


Tests for Question09 (points: 0.1)
```
>>> def test_q9():
...     try:
...         assert ans9 == 7, 'The number of items in heroes is incorrect.'
...         print('Test passed!')
...     except AssertionError as e:
...         print(e)
>>> test_q9()
Test passed!
```


Tests for Question10 (points: 0.1)
```
>>> def test_q10():
...     try:
...         assert ans10 == ['Superman', 'Spider-Man', 'Iron Man', 'Captain America', 'Bat Man', 'Aquaman', 'Ant-Man'], 'The sorted list in reverse alphabetical order is incorrect.'
...         print('Test passed!')
...     except AssertionError as e:
...         print(e)
>>> test_q10()
Test passed!
```


Tests for Question11 (points: 0.2)
```
>>> def test_q11():
...     try:
...         assert ans11 == 'IA', "The abbreviation for 'Iowa' is incorrect."
...         print('Test passed!')
...     except AssertionError as e:
...         print(e)
>>> test_q11()
Test passed!
```


Tests for Question12 (points: 0.2)
```
>>> def test_q12():
...     try:
...         assert ans12 == False, "The presence of 'Missouri' in states is incorrect."
...         print('Test passed!')
...     except AssertionError as e:
...         print(e)
>>> test_q12()
Test passed!
```


Tests for Question13 (points: 0.2)
```
>>> def test_q13():
...     try:
...         assert states['Missouri'] == 'MO', "The key-value pair for 'Missouri' was not added correctly."
...         print('Test passed!')
...     except AssertionError as e:
...         print(e)
>>> test_q13()
Test passed!
```


Tests for Question14 (points: 0.2)
```
>>> def test_q14():
...     try:
...         assert list(ans14) == ['Iowa', 'Illinois', 'Minnesota', 'Nebraska', 'Wisconsin', 'Missouri'], 'The list of keys is incorrect.'
...         print('Test passed!')
...     except AssertionError as e:
...         print(e)
>>> test_q14()
Test passed!
```


Tests for Question15 (points: 0.2)
```
>>> def test_q15():
...     try:
...         assert grades_set == {'A', 'B', 'C'}, 'The set conversion is incorrect.'
...         print('Test passed!')
...     except AssertionError as e:
...         print(e)
>>> test_q15()
Test passed!
```


Tests for Question16 (points: 0.2)
```
>>> def test_q16():
...     try:
...         assert ans16 == 3, 'The number of unique items is incorrect.'
...         print('Test passed!')
...     except AssertionError as e:
...         print(e)
>>> test_q16()
Test passed!
```


Tests for Question17 (points: 0.2)
```
>>> def test_q17():
...     try:
...         assert ans17 == True, 'The comparison of lengths is incorrect.'
...         print('Test passed!')
...     except AssertionError as e:
...         print(e)
>>> test_q17()
Test passed!
```


Tests for Question18 (points: 0.2)
```
>>> def test_q18():
...     try:
...         assert ans18 == 3, 'The remainder calculation is incorrect.'
...         print('Test passed!')
...     except AssertionError as e:
...         print(e)
>>> test_q18()
Test passed!
```


Tests for Question19 (points: 0.2)
```
>>> def test_q19():
...     try:
...         assert ans19 == 625, 'The power calculation is incorrect.'
...         print('Test passed!')
...     except AssertionError as e:
...         print(e)
>>> test_q19()
Test passed!
```


Tests for Question20 (points: 0.2)
```
>>> def test_q20():
...     try:
...         assert 'grades' not in globals(), "The variable 'grades' was not deleted."
...         print('Test passed!')
...     except AssertionError as e:
...         print(e)
>>> test_q20()
Test passed!
```
