```python
# Run this cell first! If you get any errors, uncomment the code in the next cell, then restart the kernel and try to run this cell again.
import otter
grader = otter.Notebook("Module3_Python_Basics_Part2_In-Class_Exercises.ipynb")
```

```python
# If you get any errors, uncomment the line below, run this cell, then comment out the line, restart the kernel and run the cell above. 
!pip install numpy pandas numexpr pyarrow bottleneck --upgrade
```

Output:
```
Defaulting to user installation because normal site-packages is not writeable
Requirement already satisfied: numpy in c:\users\ramachandranr\appdata\roaming\python\python312\site-packages (2.1.1)
Requirement already satisfied: pandas in c:\users\ramachandranr\appdata\roaming\python\python312\site-packages (2.2.2)
Collecting numexpr
  Downloading numexpr-2.10.1-cp312-cp312-win_amd64.whl.metadata (1.3 kB)
Collecting pyarrow
  Downloading pyarrow-17.0.0-cp312-cp312-win_amd64.whl.metadata (3.4 kB)
Collecting bottleneck
  Downloading Bottleneck-1.4.0-cp312-cp312-win_amd64.whl.metadata (8.1 kB)
Requirement already satisfied: python-dateutil>=2.8.2 in c:\users\ramachandranr\appdata\roaming\python\python312\site-packages (from pandas) (2.9.0.post0)
Requirement already satisfied: pytz>=2020.1 in c:\users\ramachandranr\appdata\roaming\python\python312\site-packages (from pandas) (2024.2)
Requirement already satisfied: tzdata>=2022.7 in c:\users\ramachandranr\appdata\roaming\python\python312\site
```

# Data Programming in Python | BAIS:6040
# Module 3 - Python Basics Part 2

## Exercises for User-Defined Functions

1\. Write a function `combine_strings` that takes two strings `str1` and `str2` as arguments and returns the concatenation of the two strings with a comma and a space between them. <i>Hint: Use the + operator to concatenate strings and include ", " between the two strings.</i>

```python
# Your answer here
 
def join_with_semicolon(lst):
    return ";".join(lst)
```

```python
# Test the function here
combine_strings("Hello", "World")  # Output: 'Hello, World'
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In[7], line 2
      1 # Test the function here
----> 2 combine_strings("Hello", "World")  # Output: 'Hello, World'

NameError: name 'combine_strings' is not defined
```

```python
grader.check("Question01")
```

2. Write a function `join_with_semicolon` that takes a list `lst` of strings as an argument and returns the concatenation of all strings in `lst` with a semicolon between them. <i>Hint: Use the join method with ";" as the separator.</i>

```python
...
```

```python
# Test the function here
join_with_semicolon(["Apple", "Banana", "Cherry"])  # Output: 'Apple;Banana;Cherry'
```

```python
grader.check("Question02")
```

3. Write a function `find_minimum` that takes two numbers `a` and `b` as arguments and returns the smaller value. If the two numbers are equal, it returns either value. <i>Hint: Use an if-else statement to compare the two numbers.</i>

```python
def find_minimum(a, b):
    # BEGIN: Function to find the minimum of two numbers
    if a < b:
        return a
    else:
        return b
    # END: Function to find the minimum of two numbers
```

```python
# Test the function here
find_minimum(7, 4)  # Output: 4
```

Output:
```
4
```

```python
grader.check("Question03")
```

Output:
```
Question03 results: All test cases passed!
```

4. Write a function `find_minimum_all` that takes a list `lst` of numbers as an argument and returns the smallest value in lst. <i>Hint: Use the min function to find the smallest value in the list.</i>

```python
def find_minimum_all(lst):
    return min(lst)
```

```python
# Test the function here
find_minimum_all([1, 3, 7, 0, 5])  # Output: 0
```

Output:
```
0
```

```python
grader.check("Question04")
```

Output:
```
Question04 results: All test cases passed!
```

5. Write a function `find_abs_minimum_all` that takes a list `lst` of numbers as an argument and returns the smallest absolute value in `lst`. <i>Hint: Use a list comprehension to create a list of absolute values, then find the minimum value.</i>


```python
def find_abs_minimum_all(lst):
    abs_lst = [abs(x) for x in lst]
    return min(abs_lst)
```

```python
# Test the function here
find_abs_minimum_all([-10, 4, -3, 2])  # Output: 2
```

Output:
```
2
```

```python
grader.check("Question05")
```

Output:
```
Question05 results: All test cases passed!
```

```python
def first_and_last(lst):
    if lst:
        return (lst[0], lst[-1])
    else:
        return (None, None)
```

```python
# Test the function here
first_and_last([1, 2, 3, 4])  # Output: (1, 4)
```

Output:
```
(1, 4)
```

```python
grader.check("Question06")
```

Output:
```
Question06 results:
    Question06 - 1 result:
        ❌ Test case failed
        Trying:
            def test_first_and_last():
                try:
                    assert first_and_last([1, 2, 3, 4]) == (1, 4)
                except AssertionError:
                    print('Hint: Use indexing to access the first and last elements of the list.')
                try:
                    assert first_and_last(['a', 'b', 'c']) == ('a', 'c')
                except AssertionError:
                    print('Hint: Ensure the function correctly handles lists of strings.')
                try:
                    assert first_and_last([True, False]) == (True, False)
                except AssertionError:
                    print('Hint: Check how the function handles lists of boolean values.')
                try:
                    assert first_and_last([1]) == 'List is too short!'
                except AssertionError:
                    print("Hint: Verify the function's behavior wi
```

7. Write a function `highlight_word` that takes two strings `text` and `word` as arguments and returns a new copy of `text` with its substring `word` highlighted in upper case. If `text` does not contain `word`, it returns “Word not found!”. <i>Hint: Use the replace method to replace the substring with its upper case version.</i>

```python
def highlight_word(text, word):
    if word in text:
        return text.replace(word, word.upper())
    else:
        return "Word not found!"
```

```python
# Test the function here
highlight_word("I love programming in Python.", "Python")  # Output: 'I love programming in PYTHON.'
```

Output:
```
'I love programming in PYTHON.'
```

```python
grader.check("Question07")
```

Output:
```
Question07 results: All test cases passed!
```

8. Write a function `slice_message` that takes a string `msg` and two integers `start_idx` and `end_idx` as arguments and returns the substring of `msg` that starts at the index position `start_idx` and ends at the index position `end_idx` - 1. <i>Hint: Use slicing to extract the substring.</i>

```python
def slice_message(msg, start_idx, end_idx):
    return msg[start_idx:end_idx]

```

```python
# Test the function here
slice_message("Data Science is fun!", 5, 12)  # Output: 'Science'
```

Output:
```
'Science'
```

```python
grader.check("Question08")
```

Output:
```
Question08 results: All test cases passed!
```

9. Write a function `slice_message2` that works the same as `slice_message` above. In addition, check inside the function if a) the type of `msg` is `str` and b) the types of `start_idx` and `end_idx` are `int`. If any of those evaluates to false, return “Invalid parameter type!”. Check also if a) `start_idx` is greater than or equal to 0, b) `start_idx` is less than `end_idx`, and c) `end_idx` is less than or equal to the length of `msg`. If any of those evaluates to false, return “Invalid parameter value!”. <i>Hint: Use isinstance to check types and conditional statements to validate values.</i>

```python
def slice_message2(msg, start_idx, end_idx):
    # Check parameter types
    if not isinstance(msg, str) or not isinstance(start_idx, int) or not isinstance(end_idx, int):
        return "Invalid parameter type!"
    
    # Check parameter values
    if start_idx < 0 or start_idx >= end_idx or end_idx > len(msg):
        return "Invalid parameter value!"
    
    return msg[start_idx:end_idx]
```

```python
# Test the function here
slice_message2("Data Science is fun!", 5, 12)  # Output: 'Science'
```

Output:
```
'Science'
```

```python
grader.check("Question09")
```

Output:
```
Question09 results: All test cases passed!
```

10. Create a dictionary `params` that sets the parameter values of `msg`, `start_idx`, and `end_idx` on your own and then call the function `slice_message2` by passing `params`. <i>Hint: Use the ** operator to unpack the dictionary when calling the function.</i>

```python
# Create a dictionary with parameter values
params = {
    "msg": "Data Science is fun!",
    "start_idx": 5,
    "end_idx": 12
}

# Call the function slice_message2 by unpacking the dictionary
result = slice_message2(**params)
print(result)  # Output: 'Science'
```

Output:
```
Science
```

```python
grader.check("Question10")
```

Output:
```
Question10 results: All test cases passed!
```

Tests for Question01 (points: None)
```
>>> def test_combine_strings():
...     try:
...         assert combine_strings('Hello', 'World') == 'Hello, World'
...     except AssertionError:
...         print('Hint: Make sure to concatenate the strings with a comma and a space between them.')
...     try:
...         assert combine_strings('Data', 'Science') == 'Data, Science'
...     except AssertionError:
...         print('Hint: Ensure the function correctly handles different input strings.')
...     try:
...         assert combine_strings('', 'Empty') == ', Empty'
...     except AssertionError:
...         print('Hint: Check how the function handles empty strings.')
...     try:
...         assert combine_strings('Single', '') == 'Single, '
...     except AssertionError:
...         print("Hint: Verify the function's behavior when the second string is empty.")
...     print('All tests passed!')
>>> test_combine_strings()
All tests passed!
```


Tests for Question02 (points: None)
```
>>> def test_join_with_semicolon():
...     try:
...         assert join_with_semicolon(['Apple', 'Banana', 'Cherry']) == 'Apple;Banana;Cherry'
...     except AssertionError:
...         print('Hint: Use the join method with a semicolon as the separator.')
...     try:
...         assert join_with_semicolon(['One', 'Two', 'Three']) == 'One;Two;Three'
...     except AssertionError:
...         print('Hint: Ensure the function correctly joins multiple strings.')
...     try:
...         assert join_with_semicolon(['']) == ''
...     except AssertionError:
...         print('Hint: Check how the function handles a list with an empty string.')
...     try:
...         assert join_with_semicolon([]) == ''
...     except AssertionError:
...         print("Hint: Verify the function's behavior when the list is empty.")
...     print('All tests passed!')
>>> test_join_with_semicolon()
All tests passed!
```


Tests for Question03 (points: None)
```
>>> def test_find_minimum():
...     try:
...         assert find_minimum(7, 4) == 4
...     except AssertionError:
...         print('Hint: Make sure the function returns the smaller of the two numbers.')
...     try:
...         assert find_minimum(3, 3) == 3
...     except AssertionError:
...         print('Hint: Ensure the function handles equal numbers correctly.')
...     try:
...         assert find_minimum(-1, 1) == -1
...     except AssertionError:
...         print('Hint: Check how the function handles negative numbers.')
...     try:
...         assert find_minimum(0, 0) == 0
...     except AssertionError:
...         print("Hint: Verify the function's behavior with zero values.")
...     print('All tests passed!')
>>> test_find_minimum()
All tests passed!
```


Tests for Question04 (points: None)
```
>>> def test_find_minimum_all():
...     try:
...         assert find_minimum_all([1, 3, 7, 0, 5]) == 0
...     except AssertionError:
...         print('Hint: Use the min function to find the smallest value in the list.')
...     try:
...         assert find_minimum_all([-1, -3, -7, 0, 5]) == -7
...     except AssertionError:
...         print('Hint: Ensure the function correctly handles negative numbers.')
...     try:
...         assert find_minimum_all([10, 20, 30]) == 10
...     except AssertionError:
...         print('Hint: Check how the function handles a list of positive numbers.')
...     try:
...         assert find_minimum_all([0]) == 0
...     except AssertionError:
...         print("Hint: Verify the function's behavior with a single-element list.")
...     print('All tests passed!')
>>> test_find_minimum_all()
All tests passed!
```


Tests for Question05 (points: None)
```
>>> def test_find_abs_minimum_all():
...     try:
...         assert find_abs_minimum_all([-10, 4, -3, 2]) == 2
...     except AssertionError:
...         print('Hint: Use a list comprehension to create a list of absolute values, then find the minimum value.')
...     try:
...         assert find_abs_minimum_all([10, -20, 30, -1]) == 1
...     except AssertionError:
...         print('Hint: Ensure the function correctly handles both positive and negative numbers.')
...     try:
...         assert find_abs_minimum_all([0, -5, 5]) == 0
...     except AssertionError:
...         print('Hint: Check how the function handles zero values.')
...     try:
...         assert find_abs_minimum_all([-100, 50, -50]) == 50
...     except AssertionError:
...         print("Hint: Verify the function's behavior with large absolute values.")
...     print('All tests passed!')
>>> test_find_abs_minimum_all()
All tests passed!
```


Tests for Question06 (points: None)
```
>>> def test_first_and_last():
...     try:
...         assert first_and_last([1, 2, 3, 4]) == (1, 4)
...     except AssertionError:
...         print('Hint: Use indexing to access the first and last elements of the list.')
...     try:
...         assert first_and_last(['a', 'b', 'c']) == ('a', 'c')
...     except AssertionError:
...         print('Hint: Ensure the function correctly handles lists of strings.')
...     try:
...         assert first_and_last([True, False]) == (True, False)
...     except AssertionError:
...         print('Hint: Check how the function handles lists of boolean values.')
...     try:
...         assert first_and_last([1]) == 'List is too short!'
...     except AssertionError:
...         print("Hint: Verify the function's behavior with a single-element list.")
...     print('All tests passed!')
>>> test_first_and_last()
All tests passed!
```


Tests for Question07 (points: None)
```
>>> def test_highlight_word():
...     try:
...         assert highlight_word('I love programming in Python.', 'Python') == 'I love programming in PYTHON.'
...     except AssertionError:
...         print('Hint: Use the replace method to replace the substring with its upper case version.')
...     try:
...         assert highlight_word('Hello World', 'World') == 'Hello WORLD'
...     except AssertionError:
...         print('Hint: Ensure the function correctly highlights the specified word.')
...     try:
...         assert highlight_word('Data Science', 'Data') == 'DATA Science'
...     except AssertionError:
...         print('Hint: Check how the function handles different input strings.')
...     try:
...         assert highlight_word('No match here', 'Python') == 'Word not found!'
...     except AssertionError:
...         print("Hint: Verify the function's behavior when the word is not found in the text.")
...     print('All tests passed!')
>>> test_highlight_word()
All tests passed!
```


Tests for Question08 (points: None)
```
>>> def test_slice_message():
...     try:
...         assert slice_message('Data Science is fun!', 5, 12) == 'Science'
...     except AssertionError:
...         print('Hint: Use slicing to extract the substring.')
...     try:
...         assert slice_message('Hello World', 0, 5) == 'Hello'
...     except AssertionError:
...         print('Hint: Ensure the function correctly handles different start and end indices.')
...     try:
...         assert slice_message('Python Programming', 7, 18) == 'Programming'
...     except AssertionError:
...         print('Hint: Check how the function handles longer strings.')
...     try:
...         assert slice_message('Short', 0, 5) == 'Short'
...     except AssertionError:
...         print("Hint: Verify the function's behavior with a string that matches the slice length.")
...     print('All tests passed!')
>>> test_slice_message()
All tests passed!
```


Tests for Question09 (points: None)
```
>>> def test_slice_message2():
...     try:
...         assert slice_message2('Data Science is fun!', 5, 12) == 'Science'
...     except AssertionError:
...         print('Hint: Use slicing to extract the substring and validate parameter types and values.')
...     try:
...         assert slice_message2('Hello World', 0, 5) == 'Hello'
...     except AssertionError:
...         print('Hint: Ensure the function correctly handles different start and end indices.')
...     try:
...         assert slice_message2('Python Programming', 7, 18) == 'Programming'
...     except AssertionError:
...         print('Hint: Check how the function handles longer strings.')
...     try:
...         assert slice_message2('Short', 0, 5) == 'Short'
...     except AssertionError:
...         print("Hint: Verify the function's behavior with a string that matches the slice length.")
...     try:
...         assert slice_message2(10000, 13, 19) == 'Invalid parameter type!'
...     except AssertionError:
...         print('Hint: Check if the function correctly identifies invalid parameter types.')
...     try:
...         assert slice_message2('Data Science is fun!', -1, 12) == 'Invalid parameter value!'
...     except AssertionError:
...         print('Hint: Ensure the function correctly identifies invalid start index values.')
...     try:
...         assert slice_message2('Data Science is fun!', 5, 5) == 'Invalid parameter value!'
...     except AssertionError:
...         print("Hint: Verify the function's behavior when start index is not less than end index.")
...     try:
...         assert slice_message2('Data Science is fun!', 5, 100) == 'Invalid parameter value!'
...     except AssertionError:
...         print('Hint: Check if the function correctly identifies end index values that exceed the string length.')
...     print('All tests passed!')
>>> test_slice_message2()
All tests passed!
```


Tests for Question10 (points: None)
```
>>> def test_params():
...     params = {'msg': 'Data Science is fun!', 'start_idx': 5, 'end_idx': 12}
...     try:
...         assert slice_message2(**params) == 'Science'
...     except AssertionError:
...         print('Hint: Ensure the dictionary parameters are correctly unpacked and passed to the function.')
...     print('All tests passed!')
>>> test_params()
All tests passed!
```
